# Программирование для всех (основы Python)

## Домашнее задание №1

*Разработчик формата сдачи заданий и фабрики автоматических тестов: Щуров И.В.<br> 
Автор задач: Тамбовцева А.А.*

**РАБОТУ ВЫПОЛНИЛ(А): впишите фамилию и имя**

**За разные задачи можно получить разное число баллов. Чтобы получить оценку 10, необходимо корректно выполнить все задачи, включая последнюю задачу со «звёздочкой».** 

Задача считается решенной, если код успешно проходит автоматические тесты и не содержит грубых нарушений («подгонка» под тесты, копирование кода другого студента). «Подгонка» под тесты – написание кода не для общего случая, а для соответствия конкретным значениям в ячейке с тестами. Пример: пишем не универсальный код для проверки чётности числа, а код, который правильно работает только на числах 2, 5, 8, потому что именно эти числа используются в `test_data`. 

**Важно:** Условия задач повышенной сложности (со «звёздочкой») включают ограничения – запрет на использование некоторых конструкций и операторов. Это не искусственное усложнение жизни, это задачи, развивающие алгоритмическое мышление. Умение выполнить задачу с минимальными затратами по времени без лишнего кода с условиями и циклами крупные компании требуют сейчас даже от аналитиков. На собеседовании вполне может встретиться задача на алгоритмы, даже если вы не разработчик, и планируете мирно анализировать и визуализировать данные.

### Формат сдачи

Для сдачи задания необходимо скачать текущий ipynb-файл, вписать в него свои решения, сохранить изменения и загрузить обновленный файл на Dropbox по предложенной на странице курса ссылке. Иметь аккаунт Dropbox для этого не требуется, достаточно пройти по ссылке и выбрать файл на компьютере.

Получить ipynb-файл можно, выбрав в Jupyter пункт меню *File → Download as... → Notebook (.ipynb)*.

### Тесты для проверки кода

На первом этапе проверки ваши решения будут проверяться с помощью автоматических тестов – ячеек с кодом, которые проверяют соответствие результатов исполнения ваших программ ожидаемым результатам в условиях задач. Эти тесты нужно использовать для проверки кода перед сдачей задания. 

Итак, для предварительной проверки задания нужно сделать следующее:

1. Скачать данный ipynb-файл на свой компьютер, открыть его в Jupyter Notebook.
2. Активировать тесты (см. ниже).
3. Вставить решение каждой задачи в ячейку для кода, следующую за его условием, там, где написан комментарий `# YOUR CODE HERE`. Отступ вашего кода должен составлять 4 пробела. Ничего не менять вокруг!
4. Запустить ячейку, в которую вы вставили код с решением. Ввести какие-то входные данные, проверить визуально правильность вывода.
5. Запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (`assertion`), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

Внимание! Если в какой-то момент забыть ввести входные данные и перейти на следующую ячейку, есть риск, что Notebook перестанет откликаться. В этом случае надо перезапустить ядро: *Kernel → Restart*. При этом потеряются все значения переменных, но сам код останется.

### Активировать тесты

Запустите следующую ячейку, чтобы иметь возможность запускать тесты.

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)
            
    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача 1. В синем море, в белой пене... (1 балл)

В Светлогорске, в музее Мирового океана есть прекрасные весы, которые позволяют узнать свой вес, измеренный в селёдках, в китах, в креветках и в других морских обитателях.

Напишите программу, которая запрашивает у пользователя с клавиатуры его вес в килограммах (целое или дробное число) и выводит на экран его вес в селёдках, округлённый до целого числа. Считайте, что средний вес селёдки равен 350 граммам.

Уточнение: считайте, что пользователь вводит допустимые значения – целые и дробные числа больше 0. Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Примеры работы программы:**

Ввод:

    56

Вывод:

    Ваш вес в селёдках: 160.

<!-- grade_id: ans01, points: 0, grade: false, solution: true -->
```python
def blue_sea():
    
    # YOUR CODE HERE

blue_sea()
```

<!-- grade_id: test01, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("56", "Ваш вес в селёдках: 160."),
    ("35", "Ваш вес в селёдках: 100."),
    ("64", "Ваш вес в селёдках: 183."),
    ("82", "Ваш вес в селёдках: 234."),
    ("40.5", "Ваш вес в селёдках: 116."),
    ("47.25", "Ваш вес в селёдках: 135.")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        blue_sea()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 2. Точность – вежливость королей (1 балл)

В машинном обучении существуют специальные меры, которые позволяют определить качество модели. Среди них — точность (*precision*), полнота (*recall*) и F-мера. F-мера вычисляется так:

$$
F = 2 \times \frac{\text{precision}\times \text{recall}}{\text{precision}+\text{recall}}
$$
 
Напишите программу, которая запрашивает с клавиатуры значение *precision*, затем значение *recall* (отдельный запрос для каждого значения), вычисляет F-меру и выводит её на экран с точностью до третьего знака после запятой.

Уточнение: считайте, что пользователь вводит допустимые значения – числа от 0 до 1 (оба числа не могут быть 0, иначе происходит деление на 0). Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Примеры работы программы:**

Ввод:

    0.9
    0.85

Вывод:

    F: 0.874

<!-- grade_id: ans02, points: 0, grade: false, solution: true -->
```python
def to_be_exact():
    
    # YOUR CODE HERE

to_be_exact()
```

<!-- grade_id: test02, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("0.9 0.85", "F: 0.874"),
    ("0.65 0.68", "F: 0.665"),
    ("0.727 0.563", "F: 0.635"), 
    ("0.41 0.5", "F: 0.451"),
    ("0 1", "F: 0.000"),
    ("1 0", "F: 0.000")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        to_be_exact()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 3. Любопытство – не порок (1 балл)

Любопытный питон Петя ловит прохожих на улице и просит назвать их какую-нибудь цифру от 0 до 9. Все ответы он записывает в виде последовательности из цифр, а затем считает сумму этих цифр. Итак, напишите программу, которая:

* сначала запрашивает число прохожих с клавиатуры (число $n$);
* затем $n$ раз просит ввести какую-нибудь цифру;
* выводит на экран полученную последовательность цифр и их сумму через пробел.

Считайте, что пользователь всегда вводит с клавиатуры корректные значения, целые положительные числа для количества прохожих и цифры от 0 до 9.

Пользоваться готовой функцией `sum()` для вычисления суммы элементов списка нельзя.

**Примеры работы программы:**

Ввод:

    4
    0
    9
    8
    2

Вывод:

    0982 19
    
Ввод:

    3
    1
    1
    5

Вывод:

    115 7
    
Ввод:

    7
    9
    9
    8
    1
    2
    3
    0

Вывод:

    9981230 32

<!-- grade_id: ans03, points: 0, grade: false, solution: true -->
```python
def curious_py():
    
    # YOUR CODE HERE

curious_py()
```

<!-- grade_id: test03, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("4 5 6 0 1", "5601 12"),
    ("7 9 9 8 1 2 3 0", "9981230 32"),
    ("2 0 1", "01 1"),
    ("3 1 1 5", "115 7"),
    ("10 3 4 7 7 3 1 0 2 3 5", "3477310235 35"),
    ("1 0", "0 0")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        curious_py()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 4. Жизнь летит, как ураган... (2 балла)

В городе Дакбург намечаются выборы. Напишите функцию `get_choices()`, которая принимает на вход строку с именами кандидатов через запятую, считает их количество и возвращает число способов упорядочить их в избирательном бюллетене.

*Подсказка.* Число способов упорядочить $n$ объектов вычисляется с помощью факториала. Факториалом $n!$ натурального числа $n$ называется произведение всех натуральных чисел от 1 до $n$ включительно. Пример:

$$
4! = 1 \times 2 \times 3 \times 4 = 24
$$

Использовать при решении готовые библиотечные функции нельзя, т.е. запрещена конструкция с `import`.

**Примеры работы программы:**

Применение функции:

    get_choices("Билли,Вилли,Дилли")

Результат:

    6
    
Применение функции:

    get_choices("Скрудж Макдак,Флинтхарт Гломгольд")

Результат:

    2
    
Применение функции:

    get_choices("Зигзаг,Дональд,Клювдия,Винт")

Результат:

    24

<!-- grade_id: ans04, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
```

<!-- grade_id: test04, points: 2, grade: true, solution: false -->
```python
assert get_choices("Скрудж Макдак,Флинтхарт Гломгольд") == 2, "Неверно"
assert get_choices("Билли,Вилли,Дилли") == 6, "Неверно"
assert get_choices("Зигзаг,Дональд,Клювдия,Винт") == 24, "Неверно"
assert get_choices("A,B,C,D,E") == 120, "Неверно"
assert get_choices("0,1,2,3,4,5,6,7,8,9") == 3628800, "Неверно"
```

### Задача 5. Кортеж Его Высочества (1 балл)

Чиполлино попросил медведя, освобождённого им из зоопарка, собрать некоторую информацию о кортеже принца Лимона, который должен проехать вечером по центральной улице, но забыл, что тот не знает, как пишутся цифры. Тогда Чиполлино придумал такую систему обозначений: 

* `&` – принц Лимон (носит плащ с длинным «хвостом»);
* `8` – министры (носят тугие ленты поверх мундира);
* `0` – придворные лимончики (без лент и плащей).

Так как порядок следования в кортеже строго регламентирован, на первом месте всегда следует принц, далее – министры, а затем уже придворные лимончики. Исходя из этого, у медведя могут получаться записи такого вида: `& 88 0000` (принц, два министра, четверо придворных), `& 8 00000`(принц, министр, пятеро придворных), и так далее.

Напишите программу, которая запрашивает с клавиатуры строку, состоящую из символов, описанных выше, и пробелов, и возвращает кортеж из этих символов и чисел, которые соответствуют количеству лимонов, занимающих разные должности.

**Примеры работы программы:**

Ввод:

    & 8 00000

Вывод:

    ('&', 1, '8', 1, '0', 5)
    
Ввод:

    & 88 0000

Вывод:

    ('&', 1, '8', 2, '0', 4)
    
Ввод:

    & 888 00

Вывод:

    ('&', 1, '8', 3, '0', 2)

<!-- grade_id: ans05, points: 0, grade: false, solution: true -->
```python
def lemongrass():
    
    # YOUR CODE HERE
    
lemongrass()
```

<!-- grade_id: test05, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("& 888 0", "('&', 1, '8', 3, '0', 1)"),
    ("& 8 00000", "('&', 1, '8', 1, '0', 5)"),
    ("& 8 000", "('&', 1, '8', 1, '0', 3)"),
    ("& 8888 000000", "('&', 1, '8', 4, '0', 6)"),
    ("& 88 000000", "('&', 1, '8', 2, '0', 6)")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        lemongrass()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 6. Мультпарад (2 балла)

Реализуйте программу, находящую мультфильмы по имени героя:

* Дан словарь `data`, в котором ключами являются имена героев, а значениями — списки названий мультфильмов, где такие герои встречаются. Гарантируется, что названия мультфильмов не повторяются, все имена и названия записаны в нижнем регистре.
* В переменную `role` считывается имя героя. Этот код уже написан. Гарантируется, что такой герой есть в словаре.
* Программа выводит названия тех, мультфильмов, где такой герой есть, но его имя не фигурирует в названии. Если подходящих названий нет — программа ничего не выводит.


**Примеры работы программы:**

Ввод:

    ежик

Вывод:

    осенние корабли
    трям! здравствуйте!
    
Ввод:

    ослик

Вывод:

    заветная мечта
    винни-пух и день забот
    
Ввод:

    енот

Вывод:

    [ничего не выводится]

```python
def mp2():

    data = {"ежик" : ["ежик в тумане", "осенние корабли", 
                      "трям! здравствуйте!", "ежик и девочка"], 
    "баба-яга" : ["баба-яга против!", "ивашка из дворца пионеров", 
                  "василиса прекрасная"], 
    "ослик" : ["как ослик грустью заболел", "козлик и ослик", 
               "заветная мечта", "винни-пух и день забот"],
    "волк" : ["серый волк энд красная шапочка", "самый маленький гном", 
             "волк и теленок", "ну, погоди"],
    "енот" : ["крошка енот"]}

    role = input()

    # YOUR CODE HERE
    
mp2()
```

<!-- grade_id: cell-aadb3a9ff88a6d9b, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("ежик", ["осенние корабли", "трям! здравствуйте!"]),
    ("баба-яга", ["ивашка из дворца пионеров", "василиса прекрасная"]),
    ("ослик", ["заветная мечта", "винни-пух и день забот"]),
    ("волк", ["самый маленький гном", "ну, погоди"]),
    ("енот", [])
]

for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        mp2()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ" 
```

### Задача 7. Города, города... (2 балла)

Дан словарь с записями, ключами которых являются числовые id, а значениями – города.

```python
cities = {100: 'Москва', 101: 'Санкт-Петербург', 
          103: 'Екатеринбург', 104: 'Пермь', 
          107: 'Красноярск', 111 : 'Киров'}
```

Напишите программу, которая просит пользователя ввести с клавиатуры сначала ключ, потом значение, а затем:

* если в словаре `cities` есть запись с таким ключом, на экран выводится сообщение "Запись с таким ключом существует. Введите другой ключ."

* если в словаре `cities` нет записи с таким ключом, в этот словарь добавляется запись с введенным ключом и значением + обновленный словарь `cities` выводится на экран.

**Примеры работы программы:**

Ввод:

    Введите ключ: 101
    Введите значение: Мурманск

Вывод:

    Запись с таким ключом существует. Введите другой ключ.

Ввод:

    Введите ключ: 117
    Введите значение: Омск

Вывод:

    {100: 'Москва', 101: 'Санкт-Петербург', 
    103: 'Екатеринбург', 104: 'Пермь', 
    107: 'Красноярск', 111 : 'Киров', 117 : 'Омск'}

Обратите внимание, id должны быть целочисленными!

```python
def goroda():
    
    cities = {100: 'Москва', 101: 'Санкт-Петербург', 
          103: 'Екатеринбург', 104: 'Пермь', 
          107: 'Красноярск', 111 : 'Киров'}
    
    # YOUR CODE HERE
    
goroda()
```

<!-- grade_id: cell-926dfd868ef4bbf8, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("100 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("101 a", "Запись с таким ключом существует. Введите другой ключ."), 
    ("103 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("104 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("107 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("111 a", "Запись с таким ключом существует. Введите другой ключ."),
    ("115 Омск", "{100: 'Москва', 101: 'Санкт-Петербург', 103: 'Екатеринбург', 104: 'Пермь', 107: 'Красноярск', 111: 'Киров', 115: 'Омск'}"),
    ("1 Казань", "{100: 'Москва', 101: 'Санкт-Петербург', 103: 'Екатеринбург', 104: 'Пермь', 107: 'Красноярск', 111: 'Киров', 1: 'Казань'}")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        goroda()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 8. Welcome to Новосибирск (2 балла)

Дан словарь, содержащий информацию по станциям метро в Новосибирске (станции в списке указаны последовательно, с севера на юг в случае Ленинской ветки, и с запада на восток в случае Дзержинской ветки).

```python
nmetro = {"Ленинская" : ["Заельцовская", "Гагаринская", "Красный проспект", 
                         "Площадь Ленина", "Октябрьская",
                           "Речной Вокзал", "Студенческая", 
                         "Площадь Маркса"], 
          "Дзержинская" : ["Площадь Гарина-Михайловского",
                           "Сибирская", "Маршала Покрышкина",
                           "Березовая Роща", "Золотая Нива"]}
```

Напишите программу, которая запрашивает у пользователя сначала название текущей станции без пробела на конце, затем направление движения (`'1'` – с севера на юг или с запада на восток, `'2'` – с юга на север или с востока на запад) и выводит на экран название следующей станции.

Считайте, что пользоватль знает, какие станции являются конечными, и не будет запрашивать станцию, следующую после конечной.

**Примеры работы программы:**

Ввод:

    Введите название текущей станции: Гагаринская
    Выберите направление движения (1 - с севера на юг или с запада на восток, 2 - с юга на север или с востока на запад): 2

Вывод:

    Следующая станция: Заельцовская


Ввод:

    Введите название текущей станции: Сибирская
    Выберите направление движения (1 - с севера на юг или с запада на восток, 2 - с юга на север или с востока на запад): 1

Вывод:

    Следующая станция: Маршала Покрышкина

```python
def metro():
    
    nmetro = {"Ленинская" : ["Заельцовская", "Гагаринская", "Красный проспект", 
                         "Площадь Ленина", "Октябрьская",
                           "Речной Вокзал", "Студенческая", 
                         "Площадь Маркса"], 
          "Дзержинская" : ["Площадь Гарина-Михайловского",
                           "Сибирская", "Маршала Покрышкина",
                           "Березовая Роща", "Золотая Нива"]}
    
    # YOUR CODE HERE

metro()
```

<!-- grade_id: cell-763eeb675f7b8c90, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("Гагаринская;2", "Следующая станция: Заельцовская"),
    ("Гагаринская;1", "Следующая станция: Красный проспект"), 
    ("Сибирская;1", "Следующая станция: Маршала Покрышкина"),
    ("Сибирская;2", "Следующая станция: Площадь Гарина-Михайловского"),
    ("Речной Вокзал;1", "Следующая станция: Студенческая"),
    ("Речной Вокзал;2", "Следующая станция: Октябрьская"),
]

for inp, out in test_data:
    with Tester(inp.split(";")) as t:
        metro()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 9. Добрые дела (2 балла)

Питон Петя решил сделать хорошее дело – раздать те книги, которые у него хранятся в нескольких экземлярах. 

Напишите функцию `to_give()`, которая:

* принимает на вход список книг `books` (список строк);
* возвращает отсортированный по алфавиту список книг, которые нужно отдать, то есть те книги, которые в списке встречаются более одного раза. 

ВАЖНО! Функция не должна ничего выводить на экран.

**Примеры работы программы:**

Применение функции:
 

    to_give(['Хобби доктора Травена', 'Хобби доктора Травена', 'Змеелов', 'Маленький принц', 
    'Питоны о Питонах. Автобиография', 'Python для детей', 'Питон для детей', 'Кролики и удавы', 
    'Маленький принц', 'Удавы и питоны. Уход и содержание', 'Питон для детей'])

Результат:

    ['Маленький принц', 'Питон для детей', 'Хобби доктора Травена']

Ввод:

    to_give(['Змеелов', 'Маленький принц', 'Хобби доктора Травена'])

Результат:

    []    

```python
# YOUR CODE HERE
```

<!-- grade_id: cell-1a9e5a2e670eb1d3, points: 2, grade: true, solution: false -->
```python
assert to_give(['Хобби доктора Травена', 
                'Хобби доктора Травена', 
                'Змеелов', 'Маленький принц', 
                'Питоны о Питонах. Автобиография', 
                'Python для детей', 'Питон для детей', 
                'Кролики и удавы', 'Маленький принц', 
                'Удавы и питоны. Уход и содержание', 
                'Питон для детей']) == ['Маленький принц', 'Питон для детей', 
                                        'Хобби доктора Травена'], 'Неверно'

assert to_give(['Змеелов', 'Маленький принц', 'Хобби доктора Травена']) == [], 'Неверно'
assert to_give(['C', 'C', 'D', 'E', 'A', 'A', 'A', 'A']) == ['A', 'C'], 'Неверно'
```

### Задача 10. Два озера – твои глаза (2 балла)

Одним прекрасным летним вечером Питон решил спеть для Анаконды песню. Он попросил своего друга подыграть ему на гитаре и запел. Всё шло хорошо, но в середине Питон понял, что в песне есть слова «И я на цыпочки привстал», задумался, как пресмыкающиеся встают на цыпочки, смутился и перестал петь (но Анаконда всё равно была в восторге).

Строки песни сохранены в список `song`, один элемент списка соответствует одной строке песни. Код для считывания строк из файла, размещённого на Github, уже написан, у всех должен работать при наличии подключения к Интернету.

Напишите программу, которая реализует описанную выше ситуацию: выводит на экран строки песни до тех пор, пока не натолкнется на строку, содержащую выражение «на цыпочки». Итоговая выдача должна выглядеть так (с многоточием в конце, это «заминка»):

    Идет по свету менестрель, поет не "тра-ля-ля",
    Он песню бережно несет про дочку короля,
    За что бродячий менестрель ей стал однажды люб,
    Пускай узнает вся земля, узнает добрый люд.
    Вот перед ней, вот перед ней, вот перед ней
    Какую трель рассыпал менестрель.
    Два озера - твои глаза, я в них окончу путь,
    Я просто выбрать не могу, в каком мне утонуть,
    Я просто выбрать не могу, и я покуда жив,
    Но я, как в песне соловей, перед тобой не лжив.
    Вот перед ней, вот перед ней, вот перед ней
    Какую трель рассыпал менестрель.
    И руки нежные твои струятся, как ручьи,
    Они текут, творя добро, они пока ничьи,
    Манят пришельца два ручья, чтоб он напился всласть
    Я только выбрать не могу, к которому припасть.
    Вот перед ней, вот перед ней, вот перед ней
    Какую трель рассыпал менестрель.
    Твои глаза, твои глаза - высоких два окна,
    И люди говорят, что в них любовь ко мне видна,
    И я на цыпочки привстал, чтоб выше стать чуть-чуть,
    ...

**Важно.** Тесты ниже проверяют только соответствие вашей выдачи ожидаемому результату. Не забывайте, что код после автоматической проверки повторно проверяется вручную, поэтому учитывайте следующее: при решении нельзя использовать готовые индексы элементов списка, то есть вручную найти номер строки, содержащей «на цыпочки», и просто вывести на экран элементы до неё. Необходимо общее решение в предположении, что список строк может выглядеть иначе, например, содержать ещё какие-то строки в самом начале.

```python
import requests
page = requests.get("https://raw.githubusercontent.com/allatambov/PyAll23/main/song.txt")
song = page.text.splitlines()
```

```python
def sing_song():
    
    # YOUR CODE HERE

sing_song()
```

<!-- grade_id: cell-96c7058bac6e1e10, points: 2, grade: true, solution: false -->
```python
test_data = [("", ['Идет по свету менестрель, поет не "тра-ля-ля",\n',
 'Он песню бережно несет про дочку короля,\n',
 'За что бродячий менестрель ей стал однажды люб,\n',
 'Пускай узнает вся земля, узнает добрый люд.\n',
 'Вот перед ней, вот перед ней, вот перед ней\n',
 'Какую трель рассыпал менестрель.\n',
 'Два озера - твои глаза, я в них окончу путь,\n',
 'Я просто выбрать не могу, в каком мне утонуть,\n',
 'Я просто выбрать не могу, и я покуда жив,\n',
 'Но я, как в песне соловей, перед тобой не лжив.\n',
 'Вот перед ней, вот перед ней, вот перед ней\n',
 'Какую трель рассыпал менестрель.\n',
 'И руки нежные твои струятся, как ручьи,\n',
 'Они текут, творя добро, они пока ничьи,\n',
 'Манят пришельца два ручья, чтоб он напился всласть\n',
 'Я только выбрать не могу, к которому припасть.\n',
 'Вот перед ней, вот перед ней, вот перед ней\n',
 'Какую трель рассыпал менестрель.\n',
 'Твои глаза, твои глаза - высоких два окна,\n',
 'И люди говорят, что в них любовь ко мне видна,\n',
 'И я на цыпочки привстал, чтоб выше стать чуть-чуть,\n',
                   '...\n'])]

for inp, out in test_data:
    with Tester([inp]) as t:
        sing_song()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверное количество строк в выводе"
        assert line_t == "".join(out), "Неверные строки"
```

### Задача 11*.  Что скрывают звёзды (3 балла)

Напишите программу, которая угадывает, какая цифра была скрыта за «звёздочкой», если известно, что пользователь вводит с клавиатуры сначала одно двузначное число со «звёздочкой» на конце, затем – второе двузначное число со «звёздочкой» в начале, а затем их сумму. Считайте, что сумма чисел тоже является двузначным числом.

При решении этой задачи нельзя использовать конструкцию `if-else` (даже если вы с ней знакомы) и любые способы явной проверки выполнения условий, например, с помощью операторов `>` или `<`.

**Примеры работы программы:**

Ввод:

    2*
    *1
    54

Вывод:

    * = 3

Ввод:

    1*
    *3
    90

Вывод:

    * = 7

<!-- grade_id: ans07, points: 0, grade: false, solution: true -->
```python
def stars():
    
    # YOUR CODE HERE

stars()
```

<!-- grade_id: test07, points: 3, grade: true, solution: false -->
```python
test_data = [
    ("2* *1 54", "* = 3"),
    ("1* *3 90", "* = 7"),
    ("2* *1 65", "* = 4"),
    ("3* *3 44", "* = 1"),
    ("1* *4 80", "* = 6"),
    ("1* *6 71", "* = 5")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        stars()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

*Названия и сюжеты задач:*

*№1: В синем море, в белой пене... (Арменфильм, 1984)* <br>
*№2: Калька с цитаты Людовика XIV* <br>
*№4: Утиные истории (Дисней, 1987-1990)*<br>
*№5: Д.Родари. Приключения Чиполлино (1951)*<br>
*№10: Песня о влюбленном менестреле», слова В.Коростылева, музыка М.Таривердиева, звучит в фильме «Ученик лекаря».*<br>